---
title: 'Module X'
description: 'Video exercise'
free_preview: true
---

## Introduction

```yaml
type: VideoExercise
key: 55f838b3cd
lang: r
xp: 50
skills: 1
```


---

## My fake multiple choice page

```yaml
type: MultipleChoiceExercise
key: 4e72afdc13
xp: 50
```

Test multiple choice question. Blah, blah, blah

`@possible_answers`
- Yes
- [No]
- Maybe

`@hint`
Here is a hint

`@pre_exercise_code`
```{r}

```

`@sct`
```{r}

```
