---
title: Module X
description: Missing data
---

## Missing data

```yaml
type: R Exercise
key: 55f838b3cd
lang: r
xp: 100
skills: 1
```

Follow the instructions for this module below.


`instructions`
We will be learning how to deal with missing data in a dataset.

- Watch this [video](https://www.youtube.com/watch?v=D1s_It47cBU) for a tutorial on how to clean a dataset. 

- Follow this [link](https://www.youtube.com/watch?v=D1s_It47cBU) to XXX.

`learning objectives`

By the end of this module you should be able to:

- Recognize missing data

- Clean dataset using a variety of methods


`pre_exercise_code`
```{r}
# Load datasets and packages here.
```

`@sample_code`
```{r}
# This
# is
# how
# to
# address
# missing
# data
# method 1
# method 2
```

`@solution`
```{r}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```
