---
title: Module X
description: 'Video exercise'
---

## Introduction

```yaml
type: Introduction
key: 55f838b3cd
lang: r
xp: 100
skills: 1
```

Follow the instructions for this module below.


`instructions`

- We will be using Canvas in this course. 
    - Watch this [video](https://www.youtube.com/watch?v=D1s_It47cBU) for a tutorial on how to use Canvas. 

- Follow this [link](https://www.youtube.com/watch?v=D1s_It47cBU) to read a course overview.

- The syllabus is located [here](https://www.youtube.com/watch?v=D1s_It47cBU)

- Set up a [LinkedIn](www.linkedin.com) account

- Information on how to contact me or the TA is located in our Canvas page

`learning objectives`

By the end of this module you should be able to:

- Describe what this course is about

- Locate important information on the course syllabus



`pre_exercise_code`
```{r}
# Load datasets and packages here.
```

`@sample_code`
```{r}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`
```{r}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```
