---
title: Module 2
description: 'Introduction to Technologies'
---

## Introduction to Technologies

```yaml
type: Module
key: 55f838b3cd
lang: r
xp: 100
skills: 1
```

Now that students know how the course is laid out, we start introducing the various technologies/tools for managing and analyzing data that we will use in this course. In this module, the intent is not to actually teach data management or analysis techniques. It’s just about how to access the various software we will use. We will not go into the details of how to use these.


`topics`

- R/Studio

- Google Docs

- Git

- GitHub

- Why R?
 
    - Students often ask me about the difference between R and other software commonly used for data analysis (e.g. SPSS, Stata, SAS, Python). Additionally, many students often wonder why we chose to teach this course using R, as opposed to some other program. I think that the web article, The popularity of data analysis software gives some information that is really useful for addressing both of these questions. So your first, and frankly one of the only, reading assignments in this course is this article. Please click the [button](www.google.com) below to be ready to discuss in our first lab session.

- S

- Stack Overflow

- Slack (Asana?)

- Finding answers

- Introduce help pages and resources


`learning objectives`

By the end of this module you should be able to:

- Summarize the key features of R/Studio

- Summarize the key features of Google Docs

- Summarize the key features of Git/Github

- Explain why we use the various tools we use for data management and analysis

`exercise`

Visit the [SPH website](https://sph.uth.edu/) for more resources.

[Example final project](https://sph.uth.edu/)


`resources`

[The popularity of data analysis software](http://r4stats.com/articles/popularity/)

[DataCamp: Spreadsheet Basics](https://www.datacamp.com/courses/spreadsheet-basics)

What is Git resource

What is GitHub resource

`practice`


`assignment`

- Create a Google account

- Create a GitHub account

- Create a Stack Overflow account

- Join the Slack workspace

- Points for signing up for GitHub, SO, and RStudio Community


`@sample_code`
```{r}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`
```{r}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```
